```python
# Initialize Otter
import otter
grader = otter.Notebook("lab08.ipynb")
```

# Lab 08: Simple Linear Regression

Welcome to Lab 08! In this lab you will implement a basic model, define loss functions, and minimize loss functions using numeric libraries. 

To receive credit for a lab, answer all questions correctly and submit before the deadline.

**Due Date:** 

**Collaboration Policy:** Data science is a collaborative activity. While you may talk with others about the labs, we ask that you **write your solutions individually**. If you do discuss the assignments with others **please include their names below** (it's a good way to learn your classmates' names).

**Collaborators:** 

List collaborators here.

Run the cell below.

Run the cell below

```python
import pandas as pd
import numpy as np

%matplotlib inline
import matplotlib.pyplot as plt
plt.rcParams['figure.figsize'] = (12, 7)
plt.rcParams['figure.dpi'] = 100
```

# 1. Standard Units

Standard units are a way of putting different kinds of observations on the same scale. The idea is to replace a datum by the number of standard deviations it is above the mean of the data. If a datum is above the mean, its value in standard units is positive; if it is below the mean, its value in standard units is negative. A datum that is above the mean by 2.5 times the standard deviation is 2.5 in standard units.

$$\text{datum in standard units } = \text{ number of SDs the datum is above the mean } = \frac{\text{original datum } - \text{ mean of original data}}{\text{SD of original data}}$$

When a list is transformed to standard units, the mean of the new list is zero, and the SD of the new list is one: that is what it means for a set of data to be in standard units. Standard units are dimensionless. If the original list has units, the original SD has the same units. To transform a measurement to standard units, we divide the measurement (minus the mean) by the SD, which cancels the original units.

If we know the mean and SD of the original data, we can restore a datum that is in standard units to the original units of measurement, as follows:

$$\text{original value }= (\text{value in standard units}) \times \text{ SD } + \text{ mean}$$

Let's write a function to standardized the values in a numerical column.

Run the cell below.

```python
nfl = pd.read_csv("data/scores_full.csv")
nfl.head()
```

```python
nfl.columns
```

What if we wanted to visualize the distribution of the away team first quarter scores `(AQ1)`. 

<!-- BEGIN QUESTION -->

**Question 1.** Create a histogram for the `AQ1` column.

```python
...
```

<!-- END QUESTION -->

How many scores are above average? Below average? Where is the average? It would be easier to determine this information from our histogram if we change to standard units. Standard units are a way of putting different kinds of observations on the same scale. The idea is to replace a datum by the number of standard deviations it is above the mean of the data.

**Question 2.** Write a function named `standard_units` that will take in a column from a dataframe and return the standardized values.

**Note:** The formula to standardize the values in an 1-dimensional array (or `pandas` Series) is 

$$z_i = \frac{x_i - \text{mean of data}}{\text{standard deviation of data}}.$$ 

A $z-$score measures the distance between a data point and the mean using standard deviations. $z-$scores can be positive or negative. The sign tells you whether the observation is above or below the mean.

```python
def standard_units(col):
    """
    Purpose
    -------
    Standardize the values in column
    
    Parameters
    ----------
    col: A 1 dimensional numpy array or a column from a dataframe
    
    """
    svs = ...
    return svs
standard_units(nfl.AQ1)
```

```python
grader.check("q2")
```

Now that we have a function we can call it on any column in the `nfl` dataframe.

```python
standard_units(nfl.AQ2)
```

```python
standard_units(nfl.HQ1)
```

```python
standard_units(nfl.HQ2)
```

<!-- BEGIN QUESTION -->

**Question 3.** Make a histogram of the standardized values from **Question 2.** 

**Note:** Use the bins that are provided

```
bins = [-5, -4, -3, -2, -1, 0, 1, 2, 3, 4, 5]
```

```python
bins = [-5, -4, -3, -2, -1, 0, 1, 2, 3, 4, 5]
...
```

<!-- END QUESTION -->

We could also write our own correlation function. Correlation coefficients are used to measure how strong a relationship is between two variables. There are several types of correlation coefficients, but the most popular is Pearson’s. Pearson’s correlation 

$$r=\frac{\sum\limits_{i=1}^n(x_i-\bar x)(y_i-\bar y)}{\sqrt{\sum\limits_{i=1}^n (x_i-\bar x)^2(y_i-\bar y)^2}}$$

($r$) is a correlation coefficient commonly used in linear regression.  The formula returns a value between $-1$ and $1$, where:

- $1$ indicates a strong positive relationship.

- $-1$ indicates a strong negative relationship.

- 0 indicates no relationship.

Let's look at an example.

Run the cell below.

```python
galton = pd.read_csv("data/galton.csv")
galton.head()
```

**Question 4.** Write a function named `corr_coef(, y)` that accepts two 1d arrays (or columns from a dataframe) and calculates the value of $r$

**Note:** Use Pearson's formula $$r=\frac{\sum\limits_{i=1}^n(x_i-\bar x)(y_i-\bar y)}{\sqrt{\sum\limits_{i=1}^n (x_i-\bar x)^2(y_i-\bar y)^2}}$$

```python
def corr_coef(x, y):
    """
    Purpose
    -------
    Compute Pearson's Correlation Coefficient 
    
    Parameters
    ----------
    x, y: A 1 dimensional numpy array or a column from a dataframe
    """
    xbar = ...
    ybar = ...
    numerator_sum = ...
    denominator_sumx = ...
    denominator_sumy = ...
    return numerator_sum/np.sqrt(denominator_sumx*denominator_sumy)
corr_coef(galton.midparentHeight, galton.childHeight)
```

```python
grader.check("q4")
```

# 2. Correlation

Another way to compute the correlation coefficient is to use standard units. To compute the correlation ocefficient using standard units we can use the formula 

$$r=\text{average of }(x_i\text{s} \text{ in standard units} \times y_i\text{s} \text{ in standard units})$$

**Question 5.** We can use a function within a function. Write function named `su_corr_coef(x,y)` that uses our `standard_units` function to compute the correlation coefficient. 

```python
def su_corr_coef(x, y):
    """
    Purpose
    -------
    Compute Pearson's Correlation Coefficient 
    
    Parameters
    ----------
    x, y: A 1 dimensional numpy array or a column from a dataframe in standard units
    """
    su_x = ...
    su_y = ...
    return ...
su_corr_coef(galton.midparentHeight, galton.childHeight)
```

```python
grader.check("q5")
```

The purpose for the previous exercises was to demonstrate writing user-defined functions. In reality we would use the `.corr()` method.

<!-- BEGIN QUESTION -->

**Question 6.**  Use the `.corr()` method to find the correlation coefficient between `midparentHeight` and `childHeight`. Your output should resemble the table below.

|                   |**midparentHeight**|**childHeight** |
|-------------------|-------------------|----------------|
|**midparentHeight**|1.00000|0.32095|
|**childHeight**    |0.32095|1.00000|

```python
...
```

<!-- END QUESTION -->

We should also visualize our data. Sometimes correlation coeffiecients can prodcue misleading results. To see an example let's take a look at Anscombe’s quartet. 

Anscombe's quartet comprises four data sets that have nearly identical simple descriptive statistics, yet have very different distributions and appear very different when graphed. Each dataset consists of eleven $(x,y)$ points. They were constructed in 1973 by the statistician Francis Anscombe to demonstrate both the importance of graphing data when analyzing it, and the effect of outliers and other influential observations on statistical properties [(Wikipedia)](https://en.wikipedia.org/wiki/Anscombe%27s_quartet).

```python
import seaborn as sns
```

```python
df = sns.load_dataset("anscombe")
```

```python
for i in ['I', 'II', 'III', 'IV']:
    print("Data set " + i)
    print(df.groupby('dataset').get_group(i)[['x', 'y']].agg(['mean', 'median', 'std']), "\n")
```

```python
for i in ['I', 'II', 'III', 'IV']:
    print("Data set " + i)
    print(df.groupby('dataset').get_group(i)[['x', 'y']].corr(), "\n")
```

Now plot each data set.

```python
ax = sns.lmplot(x = "x", y = "y", col  = "dataset", hue = "dataset", data = df,
           col_wrap = 2, ci = None, palette = "muted", height = 4,
           scatter_kws = {"s": 50, "alpha": 1})
ax.set(xticks = np.arange(np.min(df.x) - 2, np.max(df.x) + 2,1))
ax.set(xticklabels = [])
sns.set_style('white'); 
```

<!-- BEGIN QUESTION -->

**Question 7.** List the datasets for which the correlation coefficient may be misleading.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

# 3. Visualization

Now let's visualize the `galton` data set. 

<!-- BEGIN QUESTION -->

**Question 8.** Use the `galton` data set to create a scatter plot for the `midparentHeight` and `childHeight`.

```python
...
plt.title("Galton Dataset")
plt.xlabel("\n Mid-parent Height (in)")
plt.ylabel("Child Height (in)");
```

<!-- END QUESTION -->

**Question 9.** Use `standard_units` functions to change the `midparentHeight` and `childHeight` columns to standard units.

```python
su_midparentHeight = ...
su_childHeight = ...
```

```python
grader.check("q9")
```

<!-- BEGIN QUESTION -->

**Question 10.** Create a scatter plot for the `su_midparentHeight` and `su_childHeight`.

```python
...
plt.title("Galton Dataset in Standard Units")
plt.xlabel("\n Mid-parent Height (su)")
plt.ylabel("Child Height (su)");
```

<!-- END QUESTION -->

A benefit of standard units is that it is easy to find the mean (0) for each axis. It is also easy to compare values by their distance from the mean (standard deviation).

Suppose we want to find a linear equation for the line of best fit through the points (this is known as "simple linear regression"). Simple linear regression is a statistical method that allows us to summarize and study relationships between two continuous (numerical) variables:

- One variable, denoted $x$, is regarded as the predictor, explanatory, or independent variable.

- The other variable, denoted $y$, is regarded as the response, outcome, or dependent variable.

A line of best fit refers to a line through a scatter plot of data points that best expresses the relationship between those points. If our scatter plot is in standard units we can find the equation using the functions we defined earlier. 

**Note:** The equation of the line of best fit (in standard units) is $$y_{su}=r \cdot x_{su}$$

<!-- BEGIN QUESTION -->

**Question 11.** Plot the line of best (standard units) on the scatter plot that is in standard units.

```python
x_su =
y_su =
r_su =
plt.scatter(x = ..., y = ...)
plt.plot(..., ..., c = 'r')
plt.title("Galton Dataset in Standard Units")
plt.xlabel("\n Mid-parent Height (su)")
plt.ylabel("Child Height (su)");
```

<!-- END QUESTION -->

# 4. Linear Regression

`scikit-learn` is a widely used open-source Python library for machine learning, built on top of `NumPy` and some other packages. It provides the means for preprocessing data, implementing regression, classification, etc.

Let's go through the steps of implementing linear regression on our `galton` data set.

<!-- BEGIN QUESTION -->

**Question 12.** Import the `LinearRegression` function from `scikit-learn`.

```python
from sklearn.linear_model import ...
```

<!-- END QUESTION -->

**Question 13.** Create a model object for our linear regression model. 

```python
mp_model = 
```

```python
grader.check("q13")
```

Now we can fit the model. 

<!-- BEGIN QUESTION -->

**Question 14.** Fit a linear model on the `midparentHeight` and `childHeight`.

**Note:** The $x$ variable must be a dataframe.

```python
mp_model.fit(...)
```

<!-- END QUESTION -->



```python
type(mp_model)
```

<!-- BEGIN QUESTION -->

**Question 15.** Get the slope and $y-$intercept.

**Note:** These values are stored in the `model` object.

```python
y_int = ...
slope = ...
print("The y-intercept is", y_int)
print("The slope is", slope)
```

<!-- END QUESTION -->

Finally, we can overlay the scatter plot with the line of best fit.

<!-- BEGIN QUESTION -->

**Question 16.** Overlay the scatter plot with the line of best fit.

```python
y_int = 
slope = 
plt.scatter(x = ..., y = ...)
plt.plot(..., ...*... + ..., c = 'r');
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 17.** Repeat the steps in **Question 13**, **Question 14** and **Question 15** using the height of the mother (`mother`) as the predictor and the height of the child (`childHeight`) as the response.

```python
mm_model = ...
mm_model.fit(..., ...)
print("The y-intercept is" , ...)
print("The slope is ", ...)
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 18.** Create a scatter plot for the model in **Question 17**. Overlay the scatter plot with the line of best fit.

```python
...
```

<!-- END QUESTION -->



## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

When done exporting, download the .zip file by finding it in the file browswer on the left side of the screen, then right-click and select **Download**. You'll submit this .zip file for the assignment in Moodle to Gradescope for grading.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q13 (points: None)
```
>>> isinstance(mp_model, LinearRegression)
True
```


Tests for q2 (points: None)
```
>>> -1 < sum(standard_units(nfl.AQ1)) < 1
True
```


Tests for q4 (points: None)
```
>>> round(corr_coef(galton.midparentHeight, galton.childHeight), 3)
0.321
```


Tests for q5 (points: None)
```
>>> round(su_corr_coef(galton.midparentHeight, galton.childHeight), 3)
0.321
```


Tests for q9 (points: None)
```
>>> round(np.mean(su_midparentHeight))
0
```
```
>>> round(np.mean(su_childHeight))
0
```
