```python
# Initialize Otter
import otter
grader = otter.Notebook("lab05.ipynb")
```

## Lab 05: Examining the Therapeutic Touch

Welcome to Lab 05! After such an extensive introduction to programming for data science, we are finally moving into the section of the course where we can apply our new skills to answer real questions.  

In this lab, we'll use testing techniques that were introduced in lecture to test the idea of the therapeutic touch, the idea that some practitioner can feel and massage your human energy field.

For all problems that you must write our explanations and sentences for, you **must** provide your answer in the designated space. Also, throughout this lab please be sure to not re-assign variables throughout the notebook. For example, if you use `expected_proportion_correct` in your answer to one question, do not reassign it later on. Moreover, be mindful that functions that you write in earlier cells will need to be referenced to answer questions in later questions.

**Due Date:** 

**Collaboration Policy:** Data science is a collaborative activity. While you may talk with others about the labs, we ask that you **write your solutions individually**. If you do discuss the assignments with others **please include their names below** (it's a good way to learn your classmates' names).

**Collaborators:** 

List collaborators here. 

**Recommended Reading:**

* [Randomness](https://inferentialthinking.com/chapters/09/1/Conditional_Statements.html)
* [Iteration](https://inferentialthinking.com/chapters/09/2/Iteration.html)
* [Simulation](https://inferentialthinking.com/chapters/09/2/Iteration.html)
* [Empirical Distributions](https://inferentialthinking.com/chapters/10/1/Empirical_Distributions.html)
* [Sampling from a Population](https://inferentialthinking.com/chapters/10/2/Sampling_from_a_Population.html)
* [Empirical Distribution of a Statistic](https://inferentialthinking.com/chapters/10/3/Empirical_Distribution_of_a_Statistic.html)
* [Assessing Models](https://inferentialthinking.com/chapters/11/1/Assessing_Models.html)

First, set up the notebook by running the cell below.

```python
from datascience import *
import numpy as np

%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
from matplotlib import patches

from ipywidgets import interact, interactive, fixed
import ipywidgets as widgets
```

<!-- BEGIN QUESTION -->

## 1. What is the Therapeutic Touch

The Therapeutic Touch (TT) is the idea that everyone can feel the Human Energy Field (HEF) around individuals.  Those who practice TT have described different people's HEFs as "warm as Jell-O" and "tactile as taffy." 

TT was a popular technique used throughout the 20th century that was toted as a great way to bring balance to a person's health. Certain practitioners claim they have the ability to feel the HEF and can massage it in order to promote health and relaxation in individuals.

### Emily Rosa

[Emily Rosa](https://en.wikipedia.org/wiki/Emily_Rosa) was a 4th grade student who was very familiar with the world of TT, thanks to her parents, who were both medical practitioners and skeptics of TT.

For her 4th grade science fair project, Emily decided to test whether or not TT practitioners could truly interact with a person's HEF. She later went on to publish her work in TT, becoming the youngest person to have a research paper published in a peer reviewed medical journal.

### Emily's Experiment

Emily's experiment was clean, simple, and effective. Due to her parents' occupations in the medical field, she had wide access to people who claimed to be TT practitioners. 

Emily took 21 TT practitioners and used them for her science experiment. She would take a TT practitioner and ask them to extend their hands through a screen (which they can't see through). Emily would be on the other side and would flip a fair coin. Depending on how the coin landed, she would put out either her left hand or her right hand. The TT practitioner would then have to answer which hand Emily put out. If a pracitioner could truly interact with a person's HEF, it would be expected that they answered correctly.

Overall, through 210 samples, the practitioner picked the correct hand **44%** of the time. 

Emily's main goal here was to test whether or not the TT practicioners' guesses were random, like the flip of a coin. In most medical experiments, this is the norm. We want to test whether or not the treatment has an effect, **not** whether or not the treatment actually works. 

We will now begin to formulate this experiment in terms of the terminology we learned in this course. 

**Question 1.** Describe Emily’s model for how likely the TT practitioners are to choose the correct hand. What alternative model is her model meant to discredit? Discuss with students in your group to come to a conclusion. Check in with the instuctor if you are stuck.


_Type your answer here, replacing this text._

<!-- END QUESTION -->

**Question 2.** Remember that the practitioner got the correct answer 44% (0.44) of the time. According to Emily's model, on average, what proportion of times do we expect the practitioner to guess the correct hand? Make sure your answer is a decimal between 0 and 1. 


```python
expected_proportion_correct = ...
expected_proportion_correct
```

```python
grader.check("q2")
```

The goal now is to see if our deviation from this expected proportion of correct answers is due to something other than chance. 

**Question 3.** We usually use a statistic to help determine which model the evidence points towards. What is a statistic that we can use to compare outcomes under Emily’s model to what was observed? Assign `valid_stat` to an array of integer(s) representing test statistics that Emily can use: 

1. The difference between the expected percent correct and the actual percent correct.

2. The absolute difference between the expected percent correct and the actual percent correct.

3. The sum of the expected percent correct and the actual percent correct.

**Hint:** `valid_stat = make_array(choose one of the options listed above)`


```python
valid_stat = ...
valid_stat
```

```python
grader.check("q3")
```

<!-- BEGIN QUESTION -->

**Question 4.** Why is the statistic from **Question 3.** the best choice for comparing outcomes in Emily's experiment? How does it relate to the models you defined in **Question 1.**?


_Type your answer here, replacing this text._

<!-- END QUESTION -->

**Question 5.** Define the function `statistic` which takes in an expected proportion and an actual proportion, and returns the value of the statistic chosen in **Question 3.**. Assume that the argument takes in proportions, but  return your answer as a percentage. 

**Hint:** Remember we are asking for a **percentage**, not a proportion. 


```python
def statistic(expected_prop, actual_prop):
    """Calculate a statisitc as the difference between an
       expected proportion and an actual proportion.
    """
    ...
```

```python
grader.check("q5")
```

**Question 6.** Use your `statistic` function from **Question 5.** to calculate the observed statistic from Emily's experiment. 


```python
observed_statistic = ...
observed_statistic
```

```python
grader.check("q6")
```

*Is this observed statistic consistent with what we might see under Emily’s model?*

In order to answer this question, we must simulate the experiment as though Emily's model was correct, and calculate our statistic for every simulation.

#### **`sample_proportions`**

`sample_proportions` can be used to randomly sample from multiple categories when you know the proportion of data points that are expected to fall in each category. `sample_proportions` takes two arguments: the sample size and an array that contains the distribution of categories in the population (should sum to 1).

Consider flipping a fair coin, where the two outcomes (coin lands heads and coin lands tails) occur with an equal chance. We expect that half of all coin flips will land heads, and half of all coin flips will land tails.

Run the following cell to see the simulation of 10 flips of a fair coin. Let the first item of `coin_proportions` be the proportion of heads and the second item of `coin_proportions` be the proportion of tails.

```python
expected_proportion_heads = 0.5
expected_proportion_tails = 0.5
coin_proportions = make_array(expected_proportion_heads, expected_proportion_tails) 
ten_flips = sample_proportions(10, coin_proportions)
ten_flips
```

`sample_proportions` returns an array that is the same length as the proportion array that is passed through. It contains the proportion of each category that appears in the sample. 

In our example, the first item of `ten_flips` is the simulated proportion of heads and the second item of `ten_flips` is the simulated proportion of tails.

```python
simluated_proportion_heads = ten_flips.item(0)
simluated_proportion_tails = ten_flips.item(1)
print("In our simluation, " + str(simluated_proportion_heads) + " of flips were heads and " \
      + str(simluated_proportion_tails) + " of flips were tails.")
```

**Question 7.** To begin simulating, we should start by creating a representation of Emily's model to use for our simulation. This will be an array with two items in it. The first item should be the proportion of times, assuming that Emily’s model was correct, a TT practitioner picks the correct hand. The second item should be the proportion of times, under the same assumption, that the TT practitioner picks the incorrect hand. Assign `model_proportions` to this array. 

After this, we can simulate 210 hand choices, as Emily evaluated in real life, and find a single statistic to summarize this instance of the simulation. Use the `sample_proportions` function and assign the proportion of correct hand choices (out of 210) to `simulation_proportion_correct`. Lastly, use the `statistic` function that you wrote for Question 1.5 to assign `one_statistic`  to the value of the statistic for this one simulation.

**Hint:** `sample_proportions` usage can be found [here](http://data8.org/su19/python-reference.html).


```python
expected_proportion_correct = ...
expected_proportion_incorrect = ...
model_proportions = ...
simulation_proportion_correct = ...
one_statistic = ...
one_statistic
```

```python
grader.check("q7")
```

**Question 8.** Let's now see what the distribution of statistics is actually like under Emily's model. 

Define the function `simulation_and_statistic` to take in the `model_proportions` array from Question 1.7 and the expected proportion of times a TT practitioner would guess a hand correctly under Emily's model. The function should simulate Emily running through the experiment 210 times and return the statistic of this one simulation. 

**Hint:** This should follow the same pattern as the code you did in **Question 7.**.


```python
def simulation_and_statistic(model_proportions, expected_proportion_correct):  
    """Simulates 210 TT hand choices under Emily’s model and
       returns one statistic from the simulation.
    """
    ...
```

```python
grader.check("q8")
```

**Question 9.** Using the `simulation_and_statistics` function from Question 1.8, assign `simulated_statistics` to an array of 1000 statistics that you calculated under the assumption that Emily's model was true.


```python
num_repetitions = 1000
simulated_statistics = ...
for ... in ...:
...
```

```python
grader.check("q9")
```

Let's view the distribution of the simulated statistics under Emily's model, and visually compare where the observed statistic lies relative to the simulated statistics.

```python
t = Table().with_column('Simulated Statistics', simulated_statistics)
t.hist()
plt.scatter(observed_statistic, -0.01, color='red', s=100);
```

We can make a visual argument as to whether we believe the observed statistic is consistent with Emily’s model. Here, since larger values of the test statistic suggest the alternative model (where the chance of guessing the correct hand is something other than 50%), we can formalize our analysis by finding what proportion of simulated statistics were as large or larger than our observed test statistic (the area at or to the right of the observed test statistic). If this area is small enough, we’ll declare that the observed data are inconsistent with our simulated model.

**Question 10.** Calculate the proportion of simulated statistics greater than or equal to the observed statistic. 


```python
proportion_greater_or_equal = ...
proportion_greater_or_equal
```

```python
grader.check("q10")
```

<!-- BEGIN QUESTION -->

By convention, we often compare the proportion we just calculated to 0.05. If the proportion of simulated statistics greater than or equal to the observed statistic is sufficiently small (less than or equal to 0.05), then this is evidence against Emily's model. Otherwise, we don’t have any reason to doubt Emily’s model. 

This should help you make your own conclusions about Emily Rosa's experiment. 

Therapeutic touch fell out of use after this experiment, which was eventually accepted into one of the premier medical journals. TT practitioners responded and accused Emily and her family of tampering with the results, while some claimed that Emily's bad spiritual mood towards therapeutic touch made it difficult to read her HEF. Whatever it may be, Emily's experiment is a classic example about how anyone, with the right resources, can test anything they want.

Think to yourself and respond to each of the following questions.

**Question 11.** Is the data more consistent with Emily' model (practioners were randomly guessing)? Explain.


_Type your answer here, replacing this text._

<!-- END QUESTION -->

# 2. Submitting your work
You're done with Lab 05! All assignments in the course will be distributed as notebooks like this one, and you will submit your work by doing the following:

* Save your notebook

* Restart the kernel and run up to this cell.

* Run all the tests by running the cell containing `grader.check_all()`. Make sure they pass the way you expect them to.

* Run the cell below with the code `grader.export("lab05.ipynb")`.

* Download the file named `lab05.zip`, found in the explorer pane on the left side of the screen.

* Upload `lab05.zip` to the Lab 05 assignment to Gradescope for Grading.

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

When done exporting, download the .zip file by `SHIFT`-clicking on the file name and selecting **Save Link As**. Or, find the .zip file in the left side of the screen and right-click and select **Download**. You'll submit this .zip file for the assignment in Canvas to Gradescope for grading.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q10 (points: None)
```
>>> 0 <= proportion_greater_or_equal <= 1
True
```
```
>>> proportion_greater_or_equal*1000 == np.count_nonzero(simulated_statistics >= observed_statistic)
True
```


Tests for q2 (points: None)
```
>>> 0 <= expected_proportion_correct <= 1
True
```
```
>>> expected_proportion_correct == 0.5
True
```


Tests for q3 (points: None)
```
>>> isinstance(valid_stat, np.ndarray)
True
```
```
>>> int(sum(valid_stat))
2
```


Tests for q5 (points: None)
```
>>> int(round(statistic(.5,.5) + statistic(.4,.1),1))
30
```
```
>>> int(statistic(.4,.1) - statistic(.1,.4))
0
```


Tests for q6 (points: None)
```
>>> int(round(observed_statistic,2))
6
```


Tests for q7 (points: None)
```
>>> expected_proportion_correct == expected_proportion_incorrect
True
```
```
>>> len(np.unique(model_proportions))
1
```
```
>>> sum(model_proportions)
1.0
```
```
>>> # Make sure you take the item
>>> # from the array that corrresponds
>>> # to the proportion correct simulations
>>> isinstance(simulation_proportion_correct, float)
True
```
```
>>> 0 <= one_statistic <= 25
True
```


Tests for q8 (points: None)
```
>>> isinstance(simulation_and_statistic(model_proportions, .5), float)
True
```
```
>>> 0 <= simulation_and_statistic(model_proportions, .5) <= 25
True
```


Tests for q9 (points: None)
```
>>> len(simulated_statistics)
1000
```
```
>>> np.all(simulated_statistics <= 30)
True
```
```
>>> np.all(simulated_statistics >= 0)
True
```
```
>>> 0 <= simulation_and_statistic(model_proportions, expected_proportion_correct) <= 25
True
```
