```python
# Initialize Otter
import otter
grader = otter.Notebook("lab07.ipynb")
```

## Lab 07: Modeling and Estimation with Loss Functions

Welcome to Advanced Topics in Data Science for High School! Throughout the course you will complete assignments like this one. You can't learn technical subjects without hands-on practice, so these assignments are an important part of the course.

**Collaboration Policy:**

Collaborating on labs is more than okay -- it's encouraged! You should rarely remain stuck for more than a few minutes on questions in labs, so ask a neighbor or an instructor for help. Explaining things is beneficial, too -- the best way to solidify your knowledge of a subject is to explain it. You should **not** _just_ copy/paste someone else's code, but rather work together to gain understanding of the task you need to complete. 

**Due Date:**

## Today's Assignment 

In today's assignment, you'll learn about:

* mean loss functions

* absolute loss functions

First, set up the imports by running the cell below.

```python
import pandas as pd
import numpy as np
import seaborn as sns

%matplotlib inline
import matplotlib.pyplot as plt
```

## 1. Loading the Tips Dataset

To begin with, we load the `tips` dataset from the `seaborn` library.  The `tips` data contains records of tips, total bill, and information about the person who paid the bill.

```python
df = sns.load_dataset("tips")
df.head()
```

## 2. The Constant Model and Loss Functions

### Constant Model

In the modeling context, $y$ represents our "true observations", which are typically what we are trying to model. $\hat{y}$ (pronounced y "hat") represents our prediction for any model. In this lab, we will use the constant model, where our prediction for any input is a constant:

$$\hat{y} = \theta$$

$\theta$ is what we call a **parameter**. Our goal is to find the value of our parameter that **best fits our data**. We represent the optimal parameter with $\hat{\theta}$.

We call the constant model a **summary statistic**, as we are determining one number that best "summarizes" a set of values.

### Loss Functions

Loss functions are what we use to determine the optimal parameter for our model. A loss function is a measure of how well a model is able to predict the expected outcome. In other words, it measures the deviations of the predicted values from the observed values. In the formulations below $y$ represents the observed values and $\theta$ stands for our prediction.

- **Absolute Loss** (also known as the $L_1$ loss, pronounced "ell-two") is $L_1\left(\theta, y \right) =  \left| y -\theta \right|$

- **Squared Loss** (also known as the $L_2$ loss, pronounced "ell-two") is $L_2(\theta, y) = (y - \theta)^2$

<!-- BEGIN QUESTION -->

### Squared Loss Function

$$L\left(\theta , y \right)=\left(y- \theta \right)^2$$


**Question 1.** Based on the comments in the document string in the function below, implement the squared loss function on the constant model $\hat{y}=\theta$. 

**Note:** Your answer should not use any loops.


```python
def squared_loss(y_i, theta):
    """
    Purpose
    -------
    Calculate the squared loss of the observed data
    
    Parameters
    ----------
    y_i: An array of observed values
    theta: A constant representing a summary statistic
    
    Returns
    -------
    The squared loss between the observation and the summary statistic
    """
    ...
```

<!-- END QUESTION -->

Make sure that you verify that your function (`squared_loss`) is correct before moving on to the next question.

<!-- BEGIN QUESTION -->

**Question 2.** Let us now consider the case where `theta` equals 10. For arbitrary values of `y_is`, plot the squared loss using the function you implemented in the previous question. Don't forget to label your graph.

```python
# Our choice for theta
theta = ...

# Array of values for y_i
y_is = np.linspace(0, 20, 100)

# Plot
plt.xlabel(...)
plt.ylabel(...)
plt.plot(..., ...);
```

<!-- END QUESTION -->

You should notice that the value of $\theta$ is the mean of the `y_is` array. It is the location of vertex of a parabola that is concave up. 

Run the cell below to verify. 

```python
np.mean(y_is)
```

### Average Loss

Our main concern is how "good" or how "bad" the model's predictions are for an entire data set, not just one observation. The average loss of a model 

$$\frac{1}{n}\sum\limits_{i=1}^n L(\theta, y_i)$$ 

is a measure of how well the model "fits" the data. 

If squared loss is the loss function, then the average squared loss is referred to as mean squared error (MSE), and is of the following form 

$$\text{MSE}(\theta, y)=\frac{1}{n}\sum\limits_{i=1}^n (y_i-\theta)^2$$

If absolute loss is the loss function, then the average absolute loss is referred to as mean absolute error (MAE), and is of the following form 

$$\text{MAE}(\theta, y)=\frac{1}{n}\sum\limits_{i=1}^n |y_i-\theta|$$

where 

- $n$:      Number of data values.

- $i$:      $i$th value in a data set.

- $y_i$:    Value for $i$th datum.

- $\theta$: A constant representing a summary statistic.

Let's apply our knowledge to some real world data. In section **Loading the Tips Dataset** we loaded the `tips` dataset from the `seaborn` package.

In this section, you will try to find the best summary statistic $\hat \theta$ to represent the tips given in the array. The simple procedure you will use in this lab includes constructing the mean squared error (MSE) for the tips data and finding the value that minimizes the MSE.

**Question 3.** Make a `NumPy` array named `tips` using the `tip` column from the `df` dataframe that was loaded in section entitled **Loading the Tips Dataset**. 

```python
tips = ...
```

```python
grader.check("q3")
```

Now, we can extend the above loss functions to an entire dataset by taking the average. Let the dataset $D$ be the set of observations:

$$D = \{y_1, \ldots, y_n\}$$

where $y_i$ is the $i^{\text{th}}$ tip.

We can define the average loss over the dataset as:

$$L\left(\theta, y\right)=\frac{1}{n} \sum_{i=1}^n l(\theta, y_i)$$

<!-- BEGIN QUESTION -->

**Question 4.** Define the `mean_squared_error` function which computes the mean squared error given the data as an array and a value for `theta`. 

**Note:** Assume that `data` will be a `NumPy` array.

```python
def mean_squared_error(data, theta):
    """
    Purpose
    -------
    Calculate the mean square error
    
    Parameters
    ----------
    data:  An array of data values
    theta: A constant representing a summary statistic
    
    Returns
    -------
    The value of the mean square error
    """
    ...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 5.** Choose three values for $\theta$. Then use the `mean_squared_error` function with the `tips` data to see which of your choices for $\theta$ has the least loss.

```python
# Our choice for theta
theta_1 = ...
theta_2 = ...
theta_3 = ...
print("The loss for theta_1 is", mean_squared_error(tips, theta_1))
print("The loss for theta_1 is", mean_squared_error(tips, theta_2))
print("The loss for theta_1 is", mean_squared_error(tips, theta_3))
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 6.** Which one of your values for $\theta$ in the previous question had the least loss? Why do you think this was the case?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

In the cell below we plot the mean square error for different $\theta$ values on the `tips` dataset. You should save your mean squared error estimates in a list named `mse`.

**Note:** The `theta_values` are given. Make sure to label the axes on your plot.

```python
theta_values = np.linspace(1, 5, 100)
mse = [mean_squared_error(tips, t) for t in theta_values] 

plt.xlabel(r"Choice for $\theta$")
plt.ylabel(r"Mean Square Error")
plt.plot(theta_values, mse);
```

You should notice that the minimum value of MSE looks like it occurs somewhere around 3, the average of the `tips` dataset.

### Find the Minimizing Value for Our Tips Dataset

The cell below plots some arbitrary 4$^\text{th}$ degree polynomial function. 

```python
x_values = np.linspace(-4, 2.5, 100)

def fx(x):
    return 0.1 * x**4 + 0.2*x**3 + 0.2 * x **2 + 1 * x + 10

plt.plot(x_values, fx(x_values));
```

By looking at the plot, we see that the $x$ which minimizes the function is slightly larger than $-2$. What if we want the exact value?

[The function `minimize` from `scipy.optimize`](https://docs.scipy.org/doc/scipy/reference/generated/scipy.optimize.minimize.html) will attempt to minimize a function. Run the cell below to verify that minimize seems to get the answer correct.

**Note:** For this lab, we'll just use the `minimize` function. We'll discuss how `minimize` works later in the course.

```python
from scipy.optimize import minimize
minimize(fx, x0 = 1.1)
```

The `fun` value is the minimum value of the function. The `x` is the $x$ which minimizes the function. We can index into the object returned by `minimize` to get these values. We have to add the additional `[0]` at the end because the minimizing $x$ is returned as an array, but this is not necessarily the case for other attributes (i.e., `fun`). The reason for this is that `minimize` can also minimize multivariable functions.

The parameter `x0` that we passed to the `minimize` function is where the `minimize` function starts looking as it tries to find the minimum. For example, above, `minimize` started its search at $x = 1.1$ because that's where we told it to start. For the function above, it doesn't really matter what $x$ we start at because the function has only a single global minimum. More technically, the function is [convex](https://en.wikipedia.org/wiki/Convex_function), a property of functions that we will discuss later in the course.

Alas, `minimize` isn't perfect. For example, if we give it a function with many valleys (also known as local minima) it can get stuck. For example, consider the function below:

```python
w_values = np.linspace(-2, 10, 100)

def fw(w):
    return 0.1 * w**4 - 1.5*w**3 + 6 * w **2 - 1 * w + 10

plt.plot(w_values, fw(w_values));
```

If we start the minimization at $w = 6.5$, we'll get stuck in the local minimum at $w = 7.03$. 

**Note:** No matter what your actual variable is called in your function, the `minimize` routine still calls the starting point `x0`.

```python
minimize(fw, x0 = 6.5)['x'][0]
```

If we start the minimization at $w = 0.5$ or $w=-0.5$, we'll get stuck in the local minimum at $w = 0.08$. 

```python
minimize(fw, x0 = 0.5)['x'][0]
```

```python
minimize(fw, x0 = -0.5)['x'][0]
```

<!-- BEGIN QUESTION -->

**Question 7.** Using the `minimize` function, find the value of `theta` that minimizes the mean squared error for our `tips` dataset. In other words, you want to find the exact minimum of the plot that you generated in **Question 5**.

Assign `min_sq_scipy` to the value of $\theta$ that minimizes the MSE according to the `minimize` function.

You can't pass your `mean_squared_error` function to `minimize` because `mean_squared_error` has two variables: $\theta$ and `data`. `minimize` will get confused because it thinks it needs to minimize by picking the best $\theta$ and best `data` values. We only want it to use $\theta$.

Therefore, you need to pass a function of one variable, $\theta$, to the `minimize` function. This means you'll need to create a new function of only **one** variable $\theta$. In practice this is simple, but it can also be very tricky when you do this for the first time. Make sure to ask for help if you get stuck.


```python
def mean_squared_error_with_one_variable(theta):
    """
    Purpose
    -------
    Calculate the mean square error with one variable
    
    Parameters
    ----------
    theta: A constant representing a summary statistic 
    
    Returns
    -------
    The value of the mean square error.
    """
    ...
min_sq_scipy = minimize(..., x0 = 2.5)[...][...] 
min_sq_scipy
```

<!-- END QUESTION -->

Reflecting on the lab so far, we've now seen 3 ways to find the summary statistic $\hat \theta$ that minimizes the mean squared error:

1. Create a plot of the MSE for the given data array vs. $\theta$ and eyeball the minimizing $\hat \theta$.

2. Create a function that returns the MSE for a specific data array as a function of $\theta$ and use the scipy `minimize` function to find the exact $\hat \theta$ which minimizes this function.

3. Simply compute the `mean` of the data array.

At this point, you've hopefully convinced yourself that the `mean` of the data is the summary statistic that minimizes mean squared error.

### Absolute Loss Function 

$$L\left(\theta, y \right) = \left| y - \theta \right|$$

In this section, you will follow the exact same steps as above but for the absolute loss function.

<!-- BEGIN QUESTION -->

**Question 8.** In the cell below define the function `abs_loss` which returns the absolute loss given a value of $\theta$ and `y_i`. 

```python
def abs_loss(y_i, theta):
    """
    Purpose
    -------
    Calculate the absolute loss of the observed data
    
    Parameters
    ----------
    y_i: an observed value
    theta: A constant representing a summary statistic
    
    Returns
    -------
    The absolute loss between the observation and the summary statistic
    """
    ...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 9.** Let's now consider the case where `theta` equals 10. For arbitrary values of `y_is`, plot the absolute loss using the function you implemented in the previous question. Don't forget to label your graph.

```python
# Our choice for theta
theta = ...

# Array of values for y_i
y_is = np.linspace(0, 20, 100)

# Plot
plt.xlabel(...)
plt.ylabel(...)
plt.plot(..., ...);
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 10.** Define the `mean_absolute_error` function which computes the mean absolute error given the data as an array and a value for `theta`. 

**Note:** Assume that `data` will be a `NumPy` array.

```python
def mean_absolute_error(data, theta):
    """
    Purpose
    -------
    Calculate the mean absolute error
    
    Parameters
    ----------
    data:  An array of data values
    theta: A constant representing a summary statistic
    
    Returns
    -------
    The value of the mean absolute error
    """
    ...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 11.** In the cell below plot the mean absolute error for different $\theta$ values on the `tips` dataset. You should save your mean absolute error estimates in a list named `mae`.

**Note:** The `theta_values` are given. Make sure to label the axes on your plot.

```python
# Theta values
theta_values = np.linspace(2.7, 3.02, 100

# MAE values
mae = ... 

plt.xlabel(r"Choice for $\theta$")
plt.ylabel(r"...")
plt.plot(..., ...);;
```

<!-- END QUESTION -->

You should see that the plot looks somewhat similar the plot of the mean squared error. However, there are three key differences between the plots of the MSE and MAE.

1. The minimizing $\theta$ is different.

2. The plot for MAE increases linearly instead of quadratically as we move far away from the minimizing $\theta$.

3. The plot for MAE is piecewise linear instead of smooth. Each change in slope happens at the same $\theta$ value as a data point in our dataset.

### Find the Minimizing Value for the Tips Dataset 

_Type your answer here, replacing this text._

```python
def mean_absolute_error_with_one_variable(theta):
    """
    Purpose
    -------
    Calculate the mean square absolute with one variable
    
    Parameters
    ----------
    theta: A constant representing a summary statistic 
    
    Returns
    -------
    The value of the mean absolute error
    """
    ...
min_abs_scipy = minimize(..., x0 = 2.5)[...][...] 
min_abs_scipy
```

<!-- END QUESTION -->

Just like the MSE, there are three ways to compute the summary statistic $\hat \theta$ that minimizes the MAE:

1. Create a plot of the MAE for the given data array vs. $\theta$ and eyeball a minimizing $\hat \theta$.

2. Create a function that returns the MAE for a specific data array as a function of $\theta$ and use the scipy `minimize` function to find an exact $\hat \theta$ which minimizes this function.

3. Simply compute the median of the data array.

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q3 (points: None)
```
>>> len(tips)
244
```
```
>>> np.sum(np.round(tips))
733.0
```
```
>>> isinstance(tips, np.ndarray)
True
```
