```python
# Initialize Otter
import otter
grader = otter.Notebook("lab02.ipynb")
```

# Run the cell below

To run a code cell (i.e.; execute the python code inside a Jupyter notebook) you can click the play button on the ribbon underneath the name of the notebook. Before you begin click the "Run cell" button at the top that looks like ▶| or hold down `Shift` + `Return`.

## Lab 02: Fun with Text

Welcome to Advanced Topics in Data Science for High School! Throughout the course you will complete assignments like this one. You can't learn technical subjects without hands-on practice, so these assignments are an important part of the course.

**Collaboration Policy:**

Collaborating on labs is more than okay -- it's encouraged! You should rarely remain stuck for more than a few minutes on questions in labs, so ask a neighbor or an instructor for help. Explaining things is beneficial, too -- the best way to solidify your knowledge of a subject is to explain it. You should **not** _just_ copy/paste someone else's code, but rather work together to gain understanding of the task you need to complete. 

**Due Date:** 

## Today's Assignment 

In today's assignment, you'll review:

- `if/elif/else`

- `for` loops

- string methods

The code below will read in the novel *Pride and Prejudice* by Jane Austen from a local text file.  Our file was originally downloaded from [this repository](https://github.com/laumann/ds/tree/master/hashing/books) and the text itself comes from [Project Gutenberg](https://www.gutenberg.org/). The first 1000 characters are printed out below to check that it was read in correctly.

Run the cell below.

```python
with open('data/jane-austen-pride-prejudice.txt', 'r') as f:
    pap_lines = f.read()
print(pap_lines[:1000])
```

To open the file, we used the built-in `open()` function. The `open()` function returns a file object, which has a `read()` method for reading the content of the file. To simplify the code we used a `with` statement. In Python the `with` statement simplifies our read process to just two lines.

## 1. Basic Text Processing

Now we'll remove the first part and last part of the book that has the information from Project Gutenberg.

 We'll start with our text at the string "CHAPTER I" and end it just prior to where the "Transcriber's" note appears at the end.

**Question 1.** Use the `.find` string method to locate the starting position of the words 'CHAPTER I' in the `pap_lines` string.

```python
start_here = ...
start_here
```

```python
grader.check("q1")
```

<!-- BEGIN QUESTION -->

**Question 2.** Print the first sentence of Chapter I not including the title of the chapter or any extra white space.

```python
print("First sentence of Chapter I: \n")
...
print(...)
```

<!-- END QUESTION -->

Now we want to remove all the Transcriber's notes.

**Question 3.** Locate the starting position of the second chapter.

```python
end_here = ...
end_here
```

```python
grader.check("q3")
```

**Question 4.** Save only the text for Chapter I (not including the title CHAPTER I, or the `\n` characters at the beginning or the end) in the variable named `chapter_1`.

**Note:** The `\n` is an escape character for strings that should be replaced with a blank space. 

```python
chapter_1 = ...
chapter_1
```

```python
grader.check("q4")
```

## 2. Advacned Text Processing

### Counting the Number of Vowels, Consonants, and Other Characters

Next, we turn to using a loop and `if/then/else` logic to count the number of characters we have of different types.  We'll ignore any `'_'` characters as they just represent italic text and aren't really characters of the text.

**Question 5.** Change all the letters in the `chapter_1` string to lowercase.

```python
chapter_1_lowercase = ...
chapter_1_lowercase[0:25]
```

```python
grader.check("q5")
```

**Question 6.** Now loop through and count the number of vowels, consonants and other characters.

**Note:** To earn all the points for this question you **must** use a loop.

```python
vowels = ..
consonants = ...
spaces = ...
other = ...
    ...
print("The number of vowels is:", vow)
print("The number of consonants is:", con)
print("The number of other characters is:", other)
```

```python
grader.check("q6")
```

**Question 7.** Compute the proportion of vowels, consonants, and other characters to all the non-space characters. 

**Note:** Do not round the proportions. 

```python
prop_vowels = ...
prop_consonants = ...
prop_other = ...
print('The proportion of vowels is:', prop_vow)
print('The proportion of consonants is:', prop_con)
print('The proportion of other characters is:', prop_other)
```

```python
grader.check("q7")
```

## 3. Dictionary Practice

In this part, we'll count the number of words in a book using a for loop and a dictionary.  Use the code below to read in the book as a string.

```python
with open('data/charles-dickens-oliver-twist.txt', 'r', encoding='utf-8-sig') as f:
    cd_lines = f.read()
print(cd_lines[:450])
```

### Frequency of Words

Our goal in this part of the lab is to iterate through the words in the book and count the number of times each word occurs. We can use a dictionary to store this type of data. 

For instance, you'll end up with a key of "the" with a corresponding value being the number of times "the" occurred.

**Question 8.** Replace all the `\n` values with spaces and convert each word to lowercase.

```python
cd_lines_spaces = ...
cd_lines_spaces[:10]
```

```python
grader.check("q8")
```

We'll need to remove any punctuation (for instance, "end." would need the '.' removed). The object `string.punctuation` from the `string` module  has a list of punctuation marks.

Run the cell below to see.

```python
import string
print(string.punctuation)
```

To accomplish this task you can write a nested for loop that iterates over the values in `string.punctuation`. If the punctuation value exists in the `'word'` then you can use the `.replace()` method to replace it with an empty string `""`. 

**Note:** This will make words like didn't into didnt, that's fine - ignore that as it will still count appropriately for the most part.

**Question 9.** Split the text up via the spaces to get a list of 'words' you can iterate over.

```python
words_in_cd = cd_lines_spaces.split(' ')
print(words_in_cd[:10])
```

```python
grader.check("q9")
```

Now let's remove the punctuation.

**Question 10.** Remove all the punctuation from the words `words_in_cd` list. Store the words without punctuation in a new list named `words_wo_punctuation`.

**Note:** Do not include empty characters `''` in the `words_wo_punctuation` list.

```python
words_wo_punctuation = ...
...
print(words_wo_punctuation[:50])
```

```python
grader.check("q10")
```

Now the words are processed in a list and we are ready to count.

**Question 11.** Use `if/else` logic to check if the the word already exists in the dictionary. If it does, add one to the associated value. If not, add that key to the dictionary with a value of 1.

```python
word_count_dictionary = ...
...
print("chapter:", word_count_dictionary['chapter'])
print("the:", word_count_dictionary['the'])
print("of:", word_count_dictionary['of'])
```

```python
grader.check("q11")
```

Let's print out a few items from your word count dictionary.

<!-- BEGIN QUESTION -->

**Question 12.** Write a `for` loop to print out the first 20 key/value pairs in your dictionary.

```python
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 13.** Find the top ten words in Chapter I.

```python
...
```

<!-- END QUESTION -->



## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1 (points: None)
```
>>> start_here == 1844
True
```


Tests for q10 (points: None)
```
>>> len(words_wo_punctuation)
157187
```
```
>>> 0 == len([s for s in string.punctuation if s in words_wo_punctuation])
True
```


Tests for q11 (points: None)
```
>>> word_count_dictionary['chapter']
61
```
```
>>> word_count_dictionary['the']
9558
```
```
>>> word_count_dictionary['i']
1625
```
```
>>> word_count_dictionary['of']
3829
```


Tests for q3 (points: None)
```
>>> end_here
6345
```


Tests for q4 (points: None)
```
>>> len(chapter_1)
4483
```
```
>>> chapter_1[-1:]
'.'
```
```
>>> chapter_1[0:1]
'I'
```
```
>>> chapter_1.find("\n")
-1
```
```
>>> len([i for i in range(len(chapter_1)) if chapter_1.startswith("\n", i)])
0
```


Tests for q5 (points: None)
```
>>> chapter_1_lowercase.islower()
True
```


Tests for q6 (points: None)
```
>>> vow
1301
```
```
>>> con
2086
```
```
>>> other
1096
```
```
>>> len(chapter_1_lowercase) - (vow + con + other)
0
```


Tests for q7 (points: None)
```
>>> prop_vow
0.3617908787541713
```
```
>>> prop_con
0.5800889877641824
```
```
>>> prop_other
0.3047830923248053
```


Tests for q8 (points: None)
```
>>> '\n' not in cd_lines_spaces
True
```
```
>>> len([i for i in range(len(cd_lines_spaces)) if cd_lines_spaces.startswith(" ", i)])
164983
```
```
>>> len(cd_lines_spaces)
892126
```


Tests for q9 (points: None)
```
>>> len(words_in_cd)
164984
```
