```python
# Initialize Otter
import otter
grader = otter.Notebook("lab08.ipynb")
```

# Lab 08: Simple Linear Regression

Welcome to Advanced Topics in Data Science for High School! Throughout the course you will complete assignments like this one. You can't learn technical subjects without hands-on practice, so these assignments are an important part of the course.

**Collaboration Policy:**

Collaborating on labs is more than okay -- it's encouraged! You should rarely remain stuck for more than a few minutes on questions in labs, so ask a neighbor or an instructor for help. Explaining things is beneficial, too -- the best way to solidify your knowledge of a subject is to explain it. You should **not** _just_ copy/paste someone else's code, but rather work together to gain understanding of the task you need to complete. 

**Due Date:**

## Today's Assignment 

In today's assignment, you'll learn about different aspects of linear regression.

First, set up the imports by running the cell below.

```python
import pandas as pd
import numpy as np

%matplotlib inline
import matplotlib.pyplot as plt
plt.rcParams['figure.figsize'] = (12, 7)
plt.rcParams['figure.dpi'] = 100
```

## 1. Standard Units

Standard units are a way of putting different kinds of observations on the same scale. The idea is to replace a datum by the number of standard deviations it is above the mean of the data. If a datum is above the mean, its value in standard units is positive; if it is below the mean, its value in standard units is negative. A datum that is above the mean by 2.5 times the standard deviation is 2.5 in standard units.

$$\text{datum in standard units } = \text{ number of SDs the datum is above the mean } = \frac{\text{original datum } - \text{ mean of original data}}{\text{SD of original data}}$$

When a list is transformed to standard units, the mean of the new list is zero, and the SD of the new list is one: that is what it means for a set of data to be in standard units. Standard units are dimensionless. If the original list has units, the original SD has the same units. To transform a measurement to standard units, we divide the measurement (minus the mean) by the SD, which cancels the original units.

If we know the mean and SD of the original data, we can restore a datum that is in standard units to the original units of measurement, as follows:

$$\text{original value }= (\text{value in standard units}) \times \text{ SD } + \text{ mean}$$

Let's write a function to standardized the values in a numerical column.

Run the cell below.

```python
nfl = pd.read_csv("data/scores_full.csv")
nfl.head()
```

Let's examine all the column names.

```python
nfl.columns
```

What if we wanted to visualize the distribution of the away team first quarter scores `(AQ1)`. 

<!-- BEGIN QUESTION -->

**Question 1.** Create a histogram for the `AQ1` column.

```python
...
```

<!-- END QUESTION -->

How many scores are above average? Below average? Where is the average? It would be easier to determine this information from our histogram if we change to standard units. Standard units are a way of putting different kinds of observations on the same scale. The idea is to replace a datum by the number of standard deviations it is above the mean of the data.

**Question 2.** Write a function named `standard_units` that will take in a column from a dataframe and return the standardized values.

**Note:** The formula to standardize the values in an 1-dimensional array (or `pandas` Series) is 

$$z_i = \frac{x_i - \text{mean of data}}{\text{standard deviation of data}}.$$ 

A $z-$score measures the distance between a data point and the mean using standard deviations. $z-$scores can be positive or negative. The sign tells you whether the observation is above or below the mean.

```python
def standard_units(col):
    """
    Purpose
    -------
    Standardize the values in column
    
    Parameters
    ----------
    col: A 1 dimensional numpy array or a column from a dataframe
    
    """
    ...

standard_units(nfl.AQ1)
```

# END SOLUTIION

```python
grader.check("q2")
```

Now that we have a function we can call it on any column in the `nfl` dataframe. Choose 3 different numerical columns and standardize the values using the `standard_units` function.

```python
standard_units(...)
```

```python
standard_units(...)
```

```python
standard_units(...)
```

<!-- BEGIN QUESTION -->

**Question 3.** Make a histogram of the standardized values from **Question 2.** 

**Note:** Use the bins that are provided

```
bins = [-5, -4, -3, -2, -1, 0, 1, 2, 3, 4, 5]
```

```python
bins = [-5, -4, -3, -2, -1, 0, 1, 2, 3, 4, 5]
...
```

## 2. Correlation Coefficient

Correlation coefficients are used to measure how strong a linear relationship is between two variables. There are several types of correlation coefficients, but the most popular is Pearson’s. Pearson’s correlation 

$$r=\frac{\sum\limits_{i=1}^n(x_i-\bar x)(y_i-\bar y)}{\sqrt{\sum\limits_{i=1}^n (x_i-\bar x)^2(y_i-\bar y)^2}}$$

($r$) is a correlation coefficient commonly used in linear regression.  The formula returns a value between $-1$ and $1$, where:

- $1$ indicates a strong positive relationship.

- $-1$ indicates a strong negative relationship.

- 0 indicates no relationship.

Let's look at an example.

Run the cell below.

```python
galton = pd.read_csv("data/galton.csv")
galton.head()
```

**Question 4.** Write a function named `corr_coef(x, y)` that accepts two 1d arrays (or columns from a dataframe) and calculates the value of $r$

**Note:** Use Pearson's formula $$r=\frac{\sum\limits_{i=1}^n(x_i-\bar x)(y_i-\bar y)}{\sqrt{\sum\limits_{i=1}^n (x_i-\bar x)^2(y_i-\bar y)^2}}$$

```python
def corr_coef(x, y):
    """
    Purpose
    -------
    Compute Pearson's Correlation Coefficient 
    
    Parameters
    ----------
    x, y: A 1 dimensional numpy array or a column from a dataframe
    """
    ...

corr_coef(galton.midparentHeight, galton.childHeight)
```

# END SOLUTIION

```python
grader.check("q4")
```

### Correlation Coefficient in Standard Units

Another way to compute the correlation coefficient is to use standard units. To compute the correlation coefficient using standard units we can use the formula 

$$r=\text{average of }(x_{i\text{s}} \text{ in standard units} \times y_{i\text{s}} \text{ in standard units})$$

**Question 5.** We can use a function within a function. Write function named `su_corr_coef(x,y)` that uses our `standard_units` function to compute the correlation coefficient. 

```python
def su_corr_coef(x, y):
    """
    Purpose
    -------
    Compute Pearson's Correlation Coefficient 
    
    Parameters
    ----------
    x, y: A 1 dimensional array or a column from a dataframe
    """
    ...

su_corr_coef(galton.midparentHeight, galton.childHeight)
```

# END SOLUTIION

```python
grader.check("q5")
```

The purpose for the previous exercises was to demonstrate writing user-defined functions. In reality we would use the `.corr()` method.

<!-- BEGIN QUESTION -->

**Question 6.**  Use the `.corr()` method to find the correlation coefficient between `midparentHeight` and `childHeight`. Your output should resemble the table below.

|                   |**midparentHeight**|**childHeight** |
|-------------------|-------------------|----------------|
|**midparentHeight**|1.00000|0.32095|
|**childHeight**    |0.32095|1.00000|

```python
...
```

Sometimes correlation coeffiecients can prodcue misleading results. Thus, its a good idea to visualize our data (
if possible). To see an example let's take a look at Anscombe’s quartet. 

Anscombe's quartet comprises four data sets that have nearly identical simple descriptive statistics, yet have very different distributions and appear very different when graphed. Each dataset consists of eleven $(x,y)$ points. They were constructed in 1973 by the statistician Francis Anscombe to demonstrate both the importance of graphing data when analyzing it, and the effect of outliers and other influential observations on statistical properties [(Wikipedia)](https://en.wikipedia.org/wiki/Anscombe%27s_quartet).

```python
import seaborn as sns
```

```python
df = sns.load_dataset("anscombe")
```

```python
for i in ['I', 'II', 'III', 'IV']:
    print("Data set " + i)
    print(df.groupby('dataset').get_group(i)[['x', 'y']].agg(['mean', 'median', 'std']), "\n")
```

```python
for i in ['I', 'II', 'III', 'IV']:
    print("Data set " + i)
    print(df.groupby('dataset').get_group(i)[['x', 'y']].corr(), "\n")
```

Now plot each data set.

```python
ax = sns.lmplot(x = "x", y = "y", col  = "dataset", hue = "dataset", data = df,
           col_wrap = 2, ci = None, palette = "muted", height = 4,
           scatter_kws = {"s": 50, "alpha": 1})
ax.set(xticks = np.arange(np.min(df.x) - 2, np.max(df.x) + 2,1))
ax.set(xticklabels = [])
sns.set_style('white'); 
```

**Question 7.** List the datasets for which the correlation coefficient may be misleading. Explain your choices.

_Type your answer here, replacing this text._

## 3. Visualization

**Question 8.** Use the `galton` data set to create a scatter plot for the `midparentHeight` and `childHeight`.Now let's visualize the `galton` data set. 

```python
...
plt.title("Galton Dataset")
plt.xlabel("\n Mid-parent Height (in)")
plt.ylabel("Child Height (in)");
```

<!-- END QUESTION -->

**Question 9.** Use `standard_units` functions to change the `midparentHeight` and `childHeight` columns to standard units.

```python
su_midparentHeight = ...
su_childHeight = ...
```

```python
grader.check("q9")
```

<!-- BEGIN QUESTION -->

**Question 10.** Create a scatter plot for the `su_midparentHeight` and `su_childHeight`.

```python
...
plt.title("Galton Dataset in Standard Units")
plt.xlabel("\n Mid-parent Height (su)")
plt.ylabel("Child Height (su)");
```

A benefit of standard units is that it is easy to find the mean (0) for each axis. It is also easy to compare values by their distance from the mean (standard deviation).

Suppose we want to find a linear equation for the line of best fit through the points (this is known as "simple linear regression"). Simple linear regression is a statistical method that allows us to summarize and study relationships between two continuous (numerical) variables:

- One variable, denoted $x$, is regarded as the predictor, explanatory, or independent variable.

- The other variable, denoted $y$, is regarded as the response, outcome, or dependent variable.

A line of best fit refers to a line through a scatter plot of data points that best expresses the relationship between those points. If our scatter plot is in standard units we can find the equation using functions we defined earlier. 

**Note:** The equation of the line of best fit (in standard units) is $$y_{su}=r \cdot x_{su}$$

<!-- BEGIN QUESTION -->

**Question 11.** Plot the line of best (standard units) on the scatter plot that is in standard units.

```python
...
plt.scatter(x = ..., y = ...)
plt.plot(..., ..., c = 'r')
plt.title("Galton Dataset in Standard Units")
plt.xlabel("\n Mid-parent Height (su)")
plt.ylabel("Child Height (su)");
```

## 4. Linear Regression

`scikit-learn` is a widely used open-source Python library for machine learning, built on top of `NumPy` and some other packages. It provides the means for preprocessing data, implementing regression, classification, etc.

Let's go through the steps of implementing linear regression on our `galton` data set.

<!-- BEGIN QUESTION -->

**Question 12.** Import the `LinearRegression` function from `scikit-learn`.

```python
from sklearn.linear_model import ...
```

```python
grader.check("q12")
```

<!-- BEGIN QUESTION -->

<!-- END QUESTION -->

**Question 13.** Create a model object for our linear regression model. 

```python
mp_model = ...
```

<!-- END QUESTION -->

Now we can fit the model. 

<!-- BEGIN QUESTION -->

<!-- BEGIN QUESTION -->

**Question 14.** Fit a linear model on the `midparentHeight` and `childHeight`.

**Note:** The $x$ variable must be a dataframe.

```python
mp_model.fit(..., ...)
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 15.** Get the slope and $y-$intercept.

**Note:** These values are stored in the `model` object.

```python
mp_yint = ...
mp_slope = ...
print("The y-intercept is", mp_yint)
print("The slope is", mp_slope)
```

```python
grader.check("q15")
```

Finally, we can overlay the scatter plot with the line of best fit.

<!-- BEGIN QUESTION -->

**Question 16.** Overlay the scatter plot with the line of best fit.

```python
...
plt.scatter(x = ..., y = ...)
plt.plot(..., ...*... + ..., c = 'r');
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 17.** Repeat the steps in **Question 13**, **Question 14** and **Question 15** using the height of the mother (`mother`) as the predictor and the height of the child (`childHeight`) as the response.

```python
mm_model = ...
mm_model.fit(..., ...)
mm_yint = ...
mm_slope = ...
print("The y-intercept is" , mm_yint)
print("The slope is ", mm_slope)
```

```python
grader.check("q17")
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 18.** Create a scatter plot for the model in **Question 17**. Overlay the scatter plot with the line of best fit.

```python
...
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q12 (points: None)
```
>>> import sys
>>> 'sklearn.linear_model' in sys.modules
True
```


Tests for q15 (points: None)
```
>>> round(mp_yint, 1)
22.6
```
```
>>> round(mp_slope, 2)
0.64
```


Tests for q17 (points: None)
```
>>> mm_yint
46.58710391937156
```
```
>>> mm_slope
0.3145428276813158
```


Tests for q2 (points: None)
```
>>> np.sum(standard_units(nfl.AQ1))
0.0
```


Tests for q4 (points: None)
```
>>> 0 < corr_coef(galton.midparentHeight, galton.childHeight) < 0.33
True
```


Tests for q5 (points: None)
```
>>> 0 < su_corr_coef(galton.midparentHeight, galton.childHeight) < 0.33
True
```


Tests for q9 (points: None)
```
>>> -0.0001 < np.sum(su_midparentHeight) < 0.0001
True
```
```
>>> -0.0001 < np.sum(su_childHeight) < 0.0001
True
```
