```python
# Initialize Otter
import otter
grader = otter.Notebook("lab07.ipynb")
```

# Lab 07: Confidence Intervals

Welcome to Lab 07! This week, we'll learn about confidence intervals.

To receive credit for a lab, answer all questions correctly and submit before the deadline.

**Due Date:** 

**Collaboration Policy:** Data science is a collaborative activity. While you may talk with others about the labs, we ask that you **write your solutions individually**. If you do discuss the assignments with others **please include their names below** (it's a good way to learn your classmates' names).

**Collaborators:** 

List collaborators here. 

**Helpful Resource:**
- [Python Reference](http://data8.org/fa20/python-reference.html): Cheat sheet of helpful array and table methods used in Data 8!

**Recommended Reading:**

* [Confidence Intervals](https://inferentialthinking.com/chapters/13/3/Confidence_Intervals.html)
* [Using Confidence Intervals](https://inferentialthinking.com/chapters/13/4/Using_Confidence_Intervals.html)

**Notes:** 

- This lab has hidden tests on it. That means even though tests may say 100% passed, doesn't mean your final grade will be 100%. We will be running more tests for correctness once everyone turns in the homework.

- Throughout this lab, please be sure to not re-assign variables throughout the notebook. For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Moreover, please be sure to only put your written answers in the provided cells.

First, set up the notebook by running the cell below.

```python
from datascience import *
import numpy as np

%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
```

## 1. Thai Restaurants


Ben and Frank are trying see what the best Thai restaurant in Charlotte is. They survey 1500 UNC Charlotte students selected uniformly at random, and ask each student what Thai restaurant is the best. 

**Note:** This data is fabricated for the purposes of this homework.

The choices of Thai restaurant are Lucky House, Imm Thai, Thai Temple, and Thai Basil. After compiling the results, Ben and Frank release the following percentages from their sample:

|Thai Restaurant  | Percentage|
|:------------:|:------------:|
|Lucky House | 8% |
|Imm Thai | 52% |
|Thai Temple | 25% |
|Thai Basil | 15% |

These percentages represent a uniform random sample of the population of UNC Charlotte students. We will attempt to estimate the corresponding **parameters**, or the percentage of the votes that each restaurant will receive from the entire population (the entire population is all UNC Charlotte students). We will use confidence intervals to compute a range of values that reflects the uncertainty of our estimates.

The table `votes` contains the results of the survey. 

Run the cell below to load the data.

```python
votes = Table.read_table('data/votes.csv')
votes
```

**Question 1.** Complete the function `one_resampled_percentage` below. It should return Imm Thai's **percentage** of votes after simulating one bootstrap sample of `tbl`.

**Note:** `tbl` will always be in the same format as `votes`.


```python
def one_resampled_percentage(tbl):
    """A function to simulate Imm Thai's percentage of votes
       after simulating one bootstrap sample.
    """
    bootstrap = ...
    single_percentage = ...
    return ...

one_resampled_percentage(votes)
```

```python
grader.check("q1")
```

**Question 2.** Complete the `percentages_in_resamples` function such that it returns an array of 2500 bootstrapped estimates of the percentage of voters who will vote for Imm Thai. You should use the `one_resampled_percentage` function you wrote above.

**Note:** There are no public tests for this question, the autograder cell below will return 0.0% passed. 


```python
def percentages_in_resamples():
    """A function to calculate 2500 bootstrapped estimates
       of the percentage of voters that will vote for Imm Thai.
    """
    percentage_imm = make_array()
    repetitions = 2500
    for ...
        single_percentage = ...
        percentage_imm = ...
    return ...

percentages_in_resamples()
```

In the following cell, we run the function you just defined, `percentages_in_resamples`, and create a histogram of the calculated statistic for the 2,500 bootstrap estimates of the percentage of voters who voted for Imm Thai. Based on what the original Thai restaurant percentages were, does the graph seem reasonable? Talk to your teacher if you are unsure.

```python
resampled_percentages = percentages_in_resamples()
Table().with_column('Estimated Percentage', resampled_percentages).hist("Estimated Percentage")
```

**Question 3.** Using the array `resampled_percentages`, find the values at the two edges of the middle 95% of the bootstrapped percentage estimates (i.e. compute the lower and upper ends of the interval, named `imm_lower_bound` and `imm_upper_bound`, respectively.)


```python
imm_lower_bound = ...
imm_upper_bound = ...

# The following line of code wil print the results of your calculations.
# Don't change it, just run it.
print("Bootstrapped 95% confidence interval for the percentage of Imm Thai voters in the population: [{:f}, {:f}]".format(imm_lower_bound, imm_upper_bound))
```

```python
grader.check("q3")
```

**Question 4.** The survey results seem to indicate that Imm Thai is beating all the other Thai restaurants combined among voters. We would like to use confidence intervals to determine a range of likely values for Imm Thai's true lead over all the other restaurants combined. The calculation for Imm Thai's lead over Lucky House, Thai Temple, and Thai Basil combined is:

$$\text{Imm Thai's % of the vote} - \text{(Lucky House's % of the vote + Thai Temple's % of the vote + Thai Basil's % of the vote)}$$

Define the function `one_resampled_difference` that returns **exactly one value** of Imm Thai's percentage lead over Lucky House, Thai Temple, and Thai Basil combined from one bootstrap sample of `tbl`.

```python
def one_resampled_difference(tbl):
    """A function to calculate one value of Imm Thai's percentage
       lead over the other resteraunts from one bootstrap sample.
    """
    bootstrap = ...
    imm_percentage = ... 
    lh_percentage = ...
    tt_percentage = ...
    tb_percentage = ...
    return ... 
    
one_resampled_difference(votes)
```

```python
grader.check("q4")
```

<!-- BEGIN QUESTION -->

**Question 5.** Write a function called `leads_in_resamples` that finds 2,500 bootstrapped estimates (the result of calling `one_resampled_difference`) of Imm Thai's lead over Lucky House, Thai Temple, and Thai Basil combined. Plot a histogram of the resulting samples. 

**Note:** Imm Thai's lead can be negative.


```python
def leads_in_resamples():
    """A function to calculate 2500 bootstrapped estimates
       of Imm Thai's percentage lead over the other resteraunts.
    """
    repetitions = 2500
    leads = ...
    for ...
        bootstrap_lead = ...
        leads = ...
    return ...

sampled_leads = leads_in_resamples()
Table().with_column('Estimated Lead', sampled_leads).hist("Estimated Lead")
```

<!-- END QUESTION -->

**Question 6.** Use the simulated data from **Question 5** to compute an approximate 95% confidence interval for Imm Thai's true lead over Lucky House, Thai Temple, and Thai Basil combined.


```python
diff_lower_bound = ...
diff_upper_bound = ...

# The following line of code wil print the results of your calculations.
# Don't change it, just run it.
print("Bootstrapped 95% confidence interval for Imm Thai's true lead over Lucky House, Thai Temple, and Thai Basil combined: [{:f}, {:f}]".format(diff_lower_bound, diff_upper_bound))
```

```python
grader.check("q6")
```

## 2. Interpreting Confidence Intervals


We computed the following 95% confidence interval for the percentage of Imm Thai voters: 

$$[49.40, 54.47]$$

**Note:** Your answer may have been a bit different; that doesn't mean it was wrong.

<!-- BEGIN QUESTION -->

**Question 7.** Can we say there is a 95% probability that the interval [49.40, 54.47] contains the true percentage of the population that votes for Imm Thai as the best Charlotte Thai restaurant? Answer "yes" or "no" and explain your reasoning. 

**Note:** Ambiguous answers using language like "sometimes" or "maybe" will not receive credit.


_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 8.** The staff also created 70%, 90%, and 99% confidence intervals from the same sample, but we forgot to label which confidence interval represented which percentages! Match each confidence level (70%, 90%, 99%) with its corresponding interval in the cell below (e.g. __ % CI: [49.87, 54.0] $\rightarrow$ replace the blank with one of the three confidence levels). **Then**, explain your thought process and how you came up with your answers.

The intervals are below:

* [49.87, 54.00]
* [50.67, 53.27]
* [48.80, 55.40]


_Type your answer here, replacing this text._

<!-- END QUESTION -->

**Question 9.** Suppose we produced 5,000 new samples (each one a uniform random sample of 1,500 voters/students) from the population and created a 95% confidence interval from each one. Roughly how many of those 5,000 intervals do you expect will actually contain the true percentage of the population?

Assign your answer to `true_percentage_intervals`.


```python
true_percentage_intervals = ...
```

Recall the second bootstrap confidence interval you created, which estimated Imm Thai's lead over Lucky House, Thai Temple, and Thai Basil combined. Among
voters in the sample, Imm Thai's lead was 4%. The staff's 95% confidence interval for the true lead (in the population of all voters) was

$$[-0.80, 8.80]$$

Suppose we are interested in testing a simple yes-or-no question:

> "Is the percentage of votes for Imm Thai tied with the percentage of votes for Lucky House, Thai Temple, and Thai Basil combined?"

Our null hypothesis is that the percentages are equal, or equivalently, that Imm Thai's lead is exactly 0. Our alternative hypothesis is that Imm Thai's lead is not equal to 0.  In the questions below, don't compute any confidence interval yourself - use only the notebook's 95% confidence interval.


**Question 10.** Say we use a 5% P-value cutoff.  Do we fail to reject the null, reject the null, or are we unable to tell using our staff confidence interval?

Assign `restaurants_tied` to the number corresponding to the correct answer.

1. Fail to reject the null / Data is consistent with the null hypothesis.

2. Reject the null / Data is consistent with the alternative hypothesis.

3. Unable to tell using our staff confidence interval.

**Hint:** If you're confused, take a look at [this chapter](https://inferentialthinking.com/chapters/13/4/Using_Confidence_Intervals.html) of the textbook.


```python
restaurants_tied = ...
```

```python
grader.check("q10")
```

**Question 11.** What if, instead, we use a P-value cutoff of 1%? Do we fail to reject the null, reject the null, or are we unable to tell using our staff confidence interval?

Assign `cutoff_one_percent` to the number corresponding to the correct answer.

1. Fail to reject the null / Data is consistent with the null hypothesis.

2. Reject the null / Data is consistent with the alternative hypothesis.

3. Unable to tell using our staff confidence interval.


```python
cutoff_one_percent = ...
```

```python
grader.check("q11")
```

**Question 12.** What if we use a P-value cutoff of 10%? Do we reject, fail to reject, or are we unable to tell using our confidence interval?

Assign `cutoff_ten_percent` to the number corresponding to the correct answer.

1. Reject the null / Data is consistent with the alternative hypothesis
2. Fail to reject the null / Data is consistent with the null hypothesis
3. Unable to tell using our staff confidence interval


```python
cutoff_ten_percent = ...
```

```python
grader.check("q12")
```

# 3. Submitting your work

You're done with Lab 07! All assignments in the course will be distributed as notebooks like this one, and you will submit your work by doing the following:

* Save your notebook

* Restart the kernel and run up to this cell.

* Run all the tests by running the cell containing `grader.check_all()`. Make sure they pass the way you expect them to.

* Run the cell below with the code `grader.export("lab07.ipynb")`.

* Download the file named `lab07.zip`, found in the explorer pane on the left side of the screen.

* Upload `lab07.zip` to the Lab 07 assignment to Gradescope for Grading.

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

When done exporting, download the .zip file by finding it in the file browswer on the left side of the screen, then right-click and select **Download**. You'll submit this .zip file for the assignment in Canvas to Gradescope for grading.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1 (points: None)
```
>>> type(one_resampled_percentage(votes)) in set([float, np.float64])
True
```
```
>>> 35 <= one_resampled_percentage(votes) <= 65
True
```


Tests for q10 (points: None)
```
>>> restaurants_tied in [1,2,3]
True
```


Tests for q11 (points: None)
```
>>> cutoff_one_percent in [1,2,3]
True
```


Tests for q12 (points: None)
```
>>> cutoff_ten_percent in [1,2,3]
True
```


Tests for q2 (points: None)


Tests for q3 (points: None)
```
>>> 40 <= imm_lower_bound <= imm_upper_bound <= 60
True
```


Tests for q4 (points: None)
```
>>> type(one_resampled_difference(votes)) in set([float, np.float64])
True
```


Tests for q6 (points: None)
```
>>> -5 <= diff_lower_bound <= diff_upper_bound <= 12
True
```
