```python
# Initialize Otter
import otter
grader = otter.Notebook("lab02.ipynb")
```

# Lab 02: Fun with Text

Welcome to Lab 02! This lab is meant to help you familiarize yourself with the basics of Python. In this lab you will review `if/elif/else`, `for` loops, and string methods.

To receive credit for a lab, answer all questions correctly and submit before the deadline.

**Due Date:** 

**Collaboration Policy:** Data science is a collaborative activity. While you may talk with others about the labs, we ask that you **write your solutions individually**. If you do discuss the assignments with others **please include their names below** (it's a good way to learn your classmates' names).

**Collaborators:** 

List collaborators here.

The code below will read in the novel *Pride and Prejudice* by Jane Austen from a local text file.  Our file was originally downloaded from [this repository](https://github.com/laumann/ds/tree/master/hashing/books) and the text itself comes from [Project Gutenberg](https://www.gutenberg.org/). The first 1000 characters are printed out below to check that it was read in correctly.

Run the cell below.

```python
with open('data/jane-austen-pride-prejudice.txt', 'r') as f:
    pp_lines = f.read()
print(pp_lines[:1000])
```

To open the file, we used the built-in `open()` function. The `open()` function returns a file object, which has a `read()` method for reading the content of the file. To simplify the code we used a `with` statement. In Python the `with` statement simplifies our read process to just two lines.

# 1. Basic Manipulations

Now we'll remove the first part of the book that has the information from Project Gutenberg and the last part that has similar information.

 We'll start with our text at the string "CHAPTER I" and end it just prior to where the "Transcriber's" note appears at the end.

**Question 1.** Use the `.find` string method to locate the position of the words 'CHAPTER I' in the `pp_lines` string.

```python
start_here = ...
start_here
```

```python
grader.check("q1")
```

<!-- BEGIN QUESTION -->

**Question 2.** Print the first 250 characters of Chapter I not including the part we removed.

```python
print("First 250 characters: \n")
print(...)
```

<!-- END QUESTION -->

Now we want to remove all the Transcriber's notes.

**Question 3.** Locate the position of the beginning of the second chapter.

```python
end_here = ...
end_here
```

```python
grader.check("q3")
```

<!-- BEGIN QUESTION -->

**Question 4.** Print the last 250 characters Chapter I.

```python
print("Last 250 characters: \n")
print(...)
```

<!-- END QUESTION -->

**Question 5.** Save only the text for Chapter I (not including the title CHAPTER I, or the `\n` characters at the beginning or the end) in the variable named `chapter_1`.

**Note:** The `\n` is an escape character for strings that is replaced with the newline object.

```python
chapter_1 = ...
chapter_1
```

```python
grader.check("q5")
```

# 2. Counting the Number of Vowels, Consonants, and Other Characters

Next, we turn to using a loop and `if/then/else` logic to count the number of characters we have of different types.  We'll ignore any `'_'` characters as they just represent italic text and aren't really characters of the text.

**Question 6.** Change all the letters in the `chapter_1` string to lowercase.

```python
chapter_1_lowercase = ...
```

```python
grader.check("q6")
```

**Question 7.** Now loop through and count the number of vowels, consonants, spaces, and others (not including `'_'`) characters.

**Note:** To earn all the points for this question you **must** use a loop.

```python
vowels = ...
consonants = ...
spaces = ...
other = ...

for ...
print("The number of vowels is:", vowels)
print("The number of consonants is:", consonants)
print("The number of spaces are:", spaces)
print("The number of other characters is:", other)
```

```python
grader.check("q7")
```

<!-- BEGIN QUESTION -->

**Question 8.** Use the `.format` string method to print the output in the format below.

```
The number of vowels is: xxxx
The number of consonants is: xxxx
The number of spaces are: xxx
The number of other characters is: xxx
```

```python
...
```

<!-- END QUESTION -->

**Question 9.** Compute the proportion of vowels, consonants, and other characters to the non-space characters. 

**Note:** Do not round the proportions. 

```python
total_non_space = ...
prop_vowels = ...
prop_consonants = ...
prop_other = ...
print('The proportion of vowels is:', prop_vowels)
print('The proportion of consonants is:', prop_consonants)
print('The proportion of other characters is:', prop_other)
```

```python
grader.check("q9")
```

<!-- BEGIN QUESTION -->

**Question 10.** Use the `.format` string method to print the output in the format below.

```
The number of vowels is: 0.xxx
The number of consonants is: 0.xxx
The number of other characters is: 0.xxx
```

```python
...
```

<!-- END QUESTION -->

## 3. Dictionary Practice

In this part, we'll count the number of words in a book using a for loop and a dictionary.  Use the code below to read in the book as a string.

```python
with open('data/charles-dickens-oliver-twist.txt', 'r', encoding='utf-8-sig') as f:
    cd_lines = f.read()
print(cd_lines[:450])
```

Our goal in this part of the lab is to iterate through the words in the book and count the number of times each word occurs. We can use a dictionary to store this type of data. 

For instance, you'll end up with a key of "the" with a corresponding value being the number of times "the" occurred.

**Question 11.** Replace all the `\n` values with spaces and convert each word to lowercase.

```python
cd_lines_spaces = ...
```

```python
grader.check("q11")
```

We'll need to remove any punctuation (for instance, "end." would need the '.' removed). The object `string.punctuation` from the `string` module  has a list of punctuation marks.

Run the cell below to see.

```python
import string
print(string.punctuation)
```

To accomplish this task you can write a nested for loop that iterates over the values in `string.punctuation`. If the punctuation value exists in the `'word'` then you can use the `.replace()` method to replace it with an empty string `""`. 

**Note:** This will make words like didn't into didnt, that's fine - ignore that as it will still count appropriately for the most part.

**Question 12.** Split the text up via the spaces to get a list of 'words' you can iterate over.

```python
words_in_cd = ...
```

```python
grader.check("q12")
```

Next let's remove the punctuation.

**Question 13.** Remove all the punctuation from the words `words_in_cd` list. Store the words without punctuation in a new list named `words_wo_punctuation`.

```python
words_wo_punctuation = ...

for ...
```

```python
grader.check("q13")
```

Now the `'words'` are processed in a list and we are ready to count.

**Question 14.** Use `if/else` logic to check if the the word already exists in the dictionary. If it does, add one to the associated value. If not, add that key to the dictionary with a value of 1.

```python
word_count_dictionary = ...

for ...
```

```python
grader.check("q14")
```

Let's print out a few items from your word count dictionary.

<!-- BEGIN QUESTION -->

**Question 15.** Write a `for` loop to print out the first 20 key/value pairs in your dictionary.

```python
for ...
```

<!-- END QUESTION -->



---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

When done exporting, download the .zip file by finding it in the file browswer on the left side of the screen, then right-click and select **Download**. You'll submit this .zip file for the assignment in Canvas to Gradescope for grading.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1 (points: None)
```
>>> start_here
1844
```


Tests for q11 (points: None)
```
>>> "\n" not in cd_lines_spaces
True
```
```
>>> cd_lines_spaces.islower()
True
```


Tests for q12 (points: None)
```
>>> len(words_in_cd)
164984
```
```
>>> isinstance(words_in_cd, list)
True
```


Tests for q13 (points: None)
```
>>> len(words_wo_punctuation) 
157187
```
```
>>> isinstance(words_wo_punctuation, list)
True
```


Tests for q14 (points: None)
```
>>> isinstance(word_count_dictionary, dict)
True
```
```
>>> len(word_count_dictionary.keys())
11278
```
```
>>> sum(word_count_dictionary.values())
157187
```


Tests for q3 (points: None)
```
>>> end_here
6345
```


Tests for q5 (points: None)
```
>>> isinstance(chapter_1, str)
True
```
```
>>> len(chapter_1)
4483
```


Tests for q6 (points: None)
```
>>> len(chapter_1_lowercase)
4483
```
```
>>> chapter_1_lowercase.islower()
True
```


Tests for q7 (points: 0.5)
```
>>> vowels
1301
```
```
>>> consonants
2086
```
```
>>> spaces
775
```
```
>>> other
309
```


Tests for q9 (points: None)
```
>>> total_non_space
3708
```
```
>>> prop_vowels
0.35086299892125133
```
```
>>> prop_consonants
0.5625674217907227
```
```
>>> prop_other
0.08333333333333333
```
